<a href="https://colab.research.google.com/github/ramiltrinidad-ctrl/maida/blob/main/01_IntroductionRamilTrinidad.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Introduction to Jupyter

For this course, the lecture notebooks and assignments are Jupyter notebooks.  
For future courses, you will also submit notebooks via Jojie (Jupyterhub).   

This notebook demos how to work with Jupyter notebooks and what to expect from the homeworks.  


## Problem 1 [1 pt]

This is a manually graded problem that requires a valid Python code for the solution.   
The answer cell is the succeeding cell and contains the following code snippet:

    # YOUR CODE HERE
    raise NotImplementedError()
    
To answer,
1. Comment out or delete the `raise NotImplementedError()` statement.
2. Add in the relevant code.
   
**Do not** copy, delete or move that cell around.

<div class='alert alert-info'>
    
**Manually Graded Problem (Code)**

Complete the function below

</div>

<!-- grade_id: cell-121de2cf60e780e2, points: 1, grade: true, solution: true -->
```python
def get_fav_game():
    """Return my favorite game"""
    print("Diablo")
    ### YOUR CODE HERE

get_fav_game()

```

Output:
```
Diablo
```

## Problem 2 [1 pt]

The required answer may also be in the form of text.  
In this case, the answer cell would be a **markdown cell**.  

Just like in other types of answer cells, you should restrict your answer within the cell, and avoid adding or moving around cells.

<div class='alert alert-info'>
    
**Manually Graded Problem (Text)**

Pilot callsigns are typically assigned by their squadmates, usually for [something ridiculous about them](https://www.reddit.com/r/todayilearned/comments/17w4r9a/til_pilot_callsigns_are_assigned_to_you_by_your/).  
What is your callsign, and what is the context for it?

</div>

<!-- grade_id: cell-c3a8ac1bda65576a, points: 1, grade: true, solution: true -->
My call sign is rumpy, because I look like humpty dumpty

## Problem 3 [3 pts]

Sometimes the problem requires you to perform a task.  
In that case, there is no answer cell following the statement of the problem.  
You are instead expected to add as many cells as needed after the cell to accomplish the required task.

To add cells, use the "+" button (second from the left, between the floppy disk and scissors).  
To select the proper cell type, use the dropdown menu (same bar as before).  
Most of the time, your answer will either be a code cell (code answers) or a markdown cell (text answers).

Task problems are manually graded.

<div class='alert alert-info'>

**Manually Graded Task**

Answer the following in 3 separate markdown cells:
* You are due for execution - what is your last meal request?
* You are a professional wrestler in the WWE - what is your entrance theme?
* You are a legend of your field - who is narrating your autobiography?

</div>

Chicken Joy

Hulk Hogan

Conan O Brian

## Problem 4 [1 pt]

This would be an automatically graded problem so it is very important that you properly submit your answer.

<div class='alert alert-info'>
    
**Autograded Problem**

What is the area of a square with sides of 2 units?

</div>

```python
def area_2x2():
  area = 2 * 2
  return area
```

```python
print(area_2x2())
```

Output:
```
4
```

Autograded answers are usually followed by one or more cells containing automated tests.  
The visible tests are the `assert` statements.  
You **cannot** delete or edit them.  
**Do not** copy and paste these cells.

Note that the visible tests exist both as part of the grading, and to clarify ambiguities in the intended behavior.  

There can also be hidden tests.  
Below I show one such hidden test that would be part of the grading, but would not be shown in the actual assignment.  

<!-- grade_id: cell-d67d713acad56093, points: 1, grade: true, solution: false -->
```python
from numbers import Number
from numpy.testing import assert_almost_equal

area = area_2x2()
assert isinstance(area, Number)
### BEGIN HIDDEN TESTS
assert_almost_equal(area, 4)
### END HIDDEN TESTS
```

## Problem 5 [1 pt]

<div class='alert alert-info'>
    
**Autograded Problem (multiple tests)**

Complete the function below

</div>

<!-- grade_id: cell-9626ae57a69e5d87, points: 0, grade: false, solution: true -->
```python
def always_true(a):
    """Return True always"""
    ### YOUR CODE HERE
    return True

```

```python
print (always_true(1))
```

Output:
```
True
```

Autograded answers may have more than one tests cell.

<!-- grade_id: cell-459839ccb2fd81cb, points: 0.5, grade: true, solution: false -->
```python
assert always_true(1) is True
assert always_true(True) is True
```

The tests in a tests cell may all be hidden as well.  
In that case, leave it alone as usual.

<!-- grade_id: cell-994da6990cc81fce, points: 0.5, grade: true, solution: false -->
```python
ans = always_true(False)
### BEGIN HIDDEN TESTS
assert ans is True
assert always_true('') is True
### END HIDDEN TESTS
```

```python
print(ans)
```

Output:
```
True
```

```python
print (always_true)
```

Output:
```
<function always_true at 0x7a74c0670ea0>
```

This ends the introductory exercise.  
Save this notebook then submit it.  
**Do not** rename this notebook or the functions inside it unless directed - you will get a $0$ on the relevant portions.   

Make sure you don't submit past the deadlines.