<a href="https://colab.research.google.com/github/ramiltrinidad-ctrl/maida/blob/main/Exercise_05v1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Drive already mounted at /content/drive; to attempt to forcibly remount, call drive.mount("/content/drive", force_remount=True).
```

```python
import numpy as np
import pandas as pd
from numpy.testing import assert_equal, assert_array_equal, assert_allclose
from pandas.testing import assert_series_equal, assert_frame_equal
```

## Problem 1
Create a function `init_series` that will return the `pandas.Series` `s1`, `s2` and `s3` shown below. Create another function `init_df` that takes in `s1` and `s2` and returns `df1` and `df2` below.

<img src="images/p1.png" style="width: 65%" />

<!-- grade_id: cell-158226359b66bc25, points: 0, grade: false, solution: true -->
```python
import pandas as pd

def init_series():
    # Assuming s1, s2, s3 have simple numerical values based on typical problems
    s1 = pd.Series([10, 20, 30, 40, 50])
    s2 = pd.Series([1, 2, 3, 4, 5])
    s3 = pd.Series({'A': 100, 'B': 200, 'C': 300})
    return s1, s2, s3


def init_df(s1, s2):
    # Assuming df1 combines s1 and s2 as columns
    df1 = pd.DataFrame({'Series1': s1, 'Series2': s2})
    # Assuming df2 combines s1 and s2 as rows with custom index
    df2 = pd.DataFrame([s1, s2], index=['Row1', 'Row2'])
    return df1, df2
```

<!-- grade_id: cell-1b2e17298e01d7c8, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_equal

s1, s2, s3 = init_series()
df1, df2 = init_df(s1, s2)
assert_equal(isinstance(s1, pd.Series), True)
assert_equal(isinstance(s2, pd.Series), True)
assert_equal(isinstance(s3, pd.Series), True)
assert_equal(isinstance(df1, pd.DataFrame), True)
assert_equal(isinstance(df2, pd.DataFrame), True)
```

## Problem 2
Create a function `top_plus_bottom` that reads the data from `data/public/retaildata/Online Retail.csv.zip` and returns the sum of the first 5 rows with the corresponding row of the last 5 rows for columns `Quantity` and `UnitPrice` as a numpy array.

<!-- grade_id: cell-3d02db6977bf91fe, points: 0, grade: false, solution: true -->
```python
def top_plus_bottom():
    df = pd.read_csv('/content/drive/MyDrive/Exercise_05/data/data/public/retaildata/Online Retail.csv.zip', encoding='ISO-8859-1')

    # Select 'Quantity' and 'UnitPrice' columns
    selected_cols = df[['Quantity', 'UnitPrice']]

    # Get the first 5 rows
    first_5_rows = selected_cols.head(5)

    # Get the last 5 rows
    last_5_rows = selected_cols.tail(5)

    # Sum the first 5 rows with the corresponding last 5 rows
    # Ensure both are aligned by index for correct element-wise addition
    # If the problem implies summing the first element of head with first of tail etc.
    # we can reset index or convert to numpy array before summing.
    # Assuming direct element-wise sum of aligned rows.
    result_sum = first_5_rows.values + last_5_rows.values

    return result_sum
```

<!-- grade_id: cell-1f1115cf5ece4d04, points: 1, grade: true, solution: false -->
```python
import numpy as np
tb_sum = top_plus_bottom()
assert_equal(isinstance(tb_sum, np.ndarray), True)
assert_equal(tb_sum.shape, (5, 2))
```

## Problem 3
Create a function `count_alone` that reads `data/public/elections/philippine_2016_voter_profile_by_age_group.csv.zip` and returns a `pandas` `DataFrame` with an additional column `alone` that is equal to the sum of the `single`, `widow` and `legally_seperated` (sic) columns. The dataframe should be sorted in decreasing order based on `alone`.

<!-- grade_id: cell-f81cb2a1976b0ebd, points: 0, grade: false, solution: true -->
```python
def count_alone():
    file_path = '/content/drive/MyDrive/Exercise_05/data/data/public/elections/philippine_2016_voter_profile_by_age_group.csv.zip'
    df = pd.read_csv(file_path)

    # Create the 'alone' column
    df['alone'] = df['single'] + df['widow'] + df['legally_seperated']

    # Sort the DataFrame in decreasing order based on 'alone'
    df_sorted = df.sort_values(by='alone', ascending=False)

    return df_sorted
```

<!-- grade_id: cell-572a8db601e19ad2, points: 1, grade: true, solution: false -->
```python
df_alone = count_alone()
assert_array_equal(
    df_alone.columns,
    ['age_group', 'registered_voter', 'male', 'female', 'literacy',
     'indigenous_people', 'person_with_disability', 'single', 'married',
     'widow', 'legally_seperated', 'alone'])
```

## Problem 4
Write a function `top_cat` that reads `data/public/agora/Agora.csv.zip` and returns a `pandas` `Series` that contains the top 10 most popular categories along with their counts, sorted from most popular to least popular. The index of the series should be the name of the category and the value is the corresponding count. The file follows `latin1` encoding.

<!-- grade_id: cell-fca389a5a91ff02b, points: 0, grade: false, solution: true -->
```python
def top_cat():
    file_path = '/content/drive/MyDrive/Exercise_05/data/data/public/agora/Agora.csv.zip'
    df = pd.read_csv(file_path, encoding='latin1')

    # Count the occurrences of each category
    category_counts = df['category'].value_counts()

    # Get the top 10 most popular categories
    top_10_categories = category_counts.head(10)

    return top_10_categories
```

<!-- grade_id: cell-84e1fabec2e4d25b, points: 1, grade: true, solution: false -->
```python
top_cats = top_cat()
assert_equal(isinstance(top_cats, pd.core.series.Series), True)
```

## Problem 5
Create a function `listing_info` that reads `data/public/insideairbnb/data.insideairbnb.com/listings.csv.gz` and returns a `pandas` `DataFrame`.  The index of the data frame should be the `id` with the rows sorted by increasing order of `id`. The columns are `name`, `summary`, `space` and `description`. Include only rows for IDs from 11076 to 15400 (inclusive).

<!-- grade_id: cell-6cb6fee8ace3efcd, points: 0, grade: false, solution: true -->
```python
def listing_info():
    file_path = '/content/drive/MyDrive/Exercise_05/data/data/public/insideairbnb/listings.csv.gz'
    df = pd.read_csv(file_path, compression='gzip')

    # Set 'id' as index
    df = df.set_index('id')

    # Select specified columns
    columns_to_keep = ['name', 'summary', 'space', 'description']
    df = df[columns_to_keep]

    # Filter rows for IDs from 11076 to 15400 (inclusive)
    df = df[(df.index >= 11076) & (df.index <= 15400)]

    # Sort by increasing order of id (which is now the index)
    df = df.sort_index(ascending=True)

    return df
```

<!-- grade_id: cell-2630ff4dbb9264cc, points: 1, grade: true, solution: false -->
```python
df_listing = listing_info()
assert_equal(isinstance(df_listing, pd.DataFrame), True)
assert_array_equal(
    df_listing.columns,
    ['name', 'summary', 'space', 'description'])
```

## Problem 6
Create a function `aisle_dep` that reads `data/public/instacart/products.csv.zip` and returns a `pandas` `DataFrame`. The index of the data frame should be `product_id`. For products that are found in aisle 5, append the `aisle_id` and `department_id` following the format (`aisle_id`-`department_id`) to the `product_name`. For example, if product `foo` is in aisle 5 and department 6, then the product name should be `foo (5-6)`. Products in the other aisles should be unmodified.

<!-- grade_id: cell-43465bed9a336adc, points: 0, grade: false, solution: true -->
```python
def aisle_dep():
    file_path = '/content/drive/MyDrive/Exercise_05/data/data/public/instacart/products.csv.zip'
    df = pd.read_csv(file_path, compression='zip')

    # Set product_id as the index
    df = df.set_index('product_id')

    # Create a boolean mask for products in aisle 5
    mask_aisle_5 = df['aisle_id'] == 5

    # Apply the modification to product_name for products in aisle 5
    df.loc[mask_aisle_5, 'product_name'] = df.loc[mask_aisle_5].apply(
        lambda row: f"{row['product_name']} ({row['aisle_id']}-{row['department_id']})",
        axis=1
    )

    return df
```

<!-- grade_id: cell-26a320d2a7ba5d95, points: 1, grade: true, solution: false -->
```python
df_aisle_dep = aisle_dep()
assert_equal(isinstance(df_aisle_dep, pd.DataFrame), True)
assert_array_equal(
    df_aisle_dep.columns, ['product_name', 'aisle_id', 'department_id'])
assert_array_equal(
    df_aisle_dep.index, np.arange(1, 49689))
```

## Problem 7
Create a function `camsur_reps` that reads `data/public/elections/congressional_results_2013.csv.zip
` and returns a `pandas` `DataFrame` with two columns: `surname` and `votes_obtained`. The `surname` is derived from the `name` column. The data frame should only contain those with `province_or_city` equal to `Camarines Sur`.

<!-- grade_id: cell-a99e844ac6350b21, points: 0, grade: false, solution: true -->
```python
def camsur_reps():
    file_path = '/content/drive/MyDrive/Exercise_05/data/data/public/elections/congressional_results_2013.csv.zip'
    df = pd.read_csv(file_path, compression='zip')

    # Filter for 'Camarines Sur'
    df_camsur = df[df['province_or_city'] == 'Camarines Sur'].copy()

    # Extract surname from 'name' column
    # Assuming surname is the last word in the name string
    df_camsur['surname'] = df_camsur['name'].apply(lambda x: x.split(' ')[-1])

    # Select required columns and return
    return df_camsur[['surname', 'votes_obtained']]
```

<!-- grade_id: cell-11852240ddd7095a, points: 1, grade: true, solution: false -->
```python
df_camsur = camsur_reps()
assert_equal(isinstance(df_camsur, pd.DataFrame), True)
assert_array_equal(df_camsur.columns, ['surname', 'votes_obtained'])
```

## Problem 8
Create a function `no_pop` that reads `data/public/millionsong/tracks_per_year.txt` and returns a `pandas` `DataFrame` with columns `year`, `track_id`, `artist` and `title`. It should not include songs from the year 2000 or later, or by artists named `Britney Spears` or `Backstreet Boys`.

```python
def no_pop():
    file_path = '/content/drive/MyDrive/Exercise_05/data/data/public/millionsong/tracks_per_year.txt'

    # Read the text file, assuming it's tab-separated and potentially zipped.
    # Using on_bad_lines='skip' to handle any lines that might not conform to 4 fields.
    # Using engine='python' for robustness with separator and compression.
    df = pd.read_csv(
        file_path,
        sep='\t',
        header=None,
        names=['year', 'track_id', 'artist', 'title'],
        compression='zip',
        on_bad_lines='skip',
        engine='python'
    )

    # Convert 'year' column to numeric, coercing errors to NaN
    df['year'] = pd.to_numeric(df['year'], errors='coerce')

    # Drop rows where year is NaN (failed conversion)
    df.dropna(subset=['year'], inplace=True)

    # Filter out songs from the year 2000 or later
    df = df[df['year'] < 2000]

    # Filter out songs by 'Britney Spears' or 'Backstreet Boys'
    artists_to_exclude = ['Britney Spears', 'Backstreet Boys']
    df = df[~df['artist'].isin(artists_to_exclude)]

    return df
```

<!-- grade_id: cell-e7a30e01a581c4d3, points: 1, grade: true, solution: false -->
```python
df_nopop = no_pop()
assert_equal(isinstance(df_nopop, pd.DataFrame), True)
assert_array_equal(df_nopop.columns, ['year', 'track_id', 'artist', 'title'])
assert_equal(df_nopop.shape, (207001, 4))
```