<a href="https://colab.research.google.com/github/ramiltrinidad-ctrl/maida/blob/main/Exercise_04v1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
from numpy.testing import (assert_equal, assert_almost_equal, assert_raises,
                           assert_raises_regex)
```

```python
import numpy as np
```

For all problems, make sure to put in a complete docstring if there is none yet. You may only use the Python Standard Library.

## Problem 1
Write a function `value_counts` that accepts a list `a_list` and an output filepath `out_path`. It saves the number of times each element in `a_list` appears to `out_path` as a comma-separated separated values file. Sort them from most to least frequent then by element if they have the same count. Don't include a header.

```python

```

<!-- grade_id: cell-612f8d5c295e6aaa, points: 0, grade: false, solution: true -->
```python
import collections

def value_counts(a_list, out_path):
    """
    Calculates the frequency of each element in a list and saves it to a CSV file.

    Parameters
    ----------
    a_list : list
        The input list.
    out_path : str
        The path to the output CSV file.
    """
    counts = collections.Counter(a_list)
    # Sort by count (descending), then by element (ascending)
    sorted_counts = sorted(counts.items(), key=lambda item: (-item[1], item[0]))

    with open(out_path, 'w') as f:
        for item, count in sorted_counts:
            f.write(f"{item},{count}\n")
```

<!-- grade_id: cell-6741514de37183fb, points: 1, grade: true, solution: false -->
```python
!rm -f out_file.csv
value_counts(['a', 'b', 'c'], 'out_file.csv')
with open('out_file.csv') as f:
    assert_equal(
        f.read(),
        'a,1\n'
        'b,1\n'
        'c,1\n'
    )
!rm -f out_file.csv
value_counts(['a', 'b', 'c', 'b', 'a'], 'out_file.csv')
with open('out_file.csv') as f:
    assert_equal(
        f.read(),
        'a,2\n'
        'b,2\n'
        'c,1\n'
    )
!rm -f out_file.csv
```

## Problem 2
Create a function `is_subset` that checks whether `sublist` is a subset of `superlist`.

<!-- grade_id: cell-915b3dbfa6d30678, points: 0, grade: false, solution: true -->
```python
import collections

def is_subset(sublist, superlist, strict=True):
    """
    Check whether `sublist` is a subset of `superlist`

    Parameters
    ----------
    sublist : list
        List to check whether it is a subset of `superlist`
    superlist : list
        List to check whether `sublist` is one of its subsets
    strict : bool
        If `True`, the exact sequence of `sublist` must be found in
        `superlist` for the `sublist` to be considered as a subset. If
        `False`, `sublist` will be considered a subset of `superlist` as long
        as all members of `sublist` are found in `superlist`

    Returns
    -------
    is_subset : bool
        `True` is `sublist` is a subset of `superlist`, `False` otherwise
    """
    if strict:
        if not sublist:  # An empty list is a subset of any list in strict mode
            return True
        n_sub = len(sublist)
        n_super = len(superlist)
        if n_sub > n_super:
            return False
        for i in range(n_super - n_sub + 1):
            if superlist[i:i + n_sub] == sublist:
                return True
        return False
    else:
        # Non-strict: check if all elements in sublist are present in superlist,
        # accounting for duplicates.
        sub_counts = collections.Counter(sublist)
        super_counts = collections.Counter(superlist)

        for item, count in sub_counts.items():
            if super_counts[item] < count:
                return False
        return True
```

<!-- grade_id: cell-637e733c312fd315, points: 1, grade: true, solution: false -->
```python
assert_equal(is_subset([1, 2, 3], [0, 1, 2, 3, 4, 5]), True)
assert_equal(is_subset([1, 2, 3], [0, 1, 2, 4, 3, 5]), False)
assert_equal(is_subset([1, 2, 3], [0, 1, 2, 4, 3, 5], False), True)
assert_equal(is_subset([1, 'a', 3], [0, 1, 'a', 3, 4, 5]), True)
assert_equal(is_subset([1, 'a', 3], [0, 1, 'a', 4, 3, 5]), False)
assert_equal(is_subset([1, 'a', 3], [0, 1, 'a', 4, 3, 5], False), True)
```

## Problem 3
Write a function called `has_duplicates` that takes a list and returns `True` if there is any element that appears more than once, `False` otherwise.

<!-- grade_id: cell-ada56c95461c367f, points: 0, grade: false, solution: true -->
```python
def has_duplicates(a_list):
    """
    Checks if a list contains duplicate elements.

    Parameters
    ----------
    a_list : list
        The input list.

    Returns
    -------
    bool
        True if there are duplicates, False otherwise.
    """
    return len(a_list) != len(set(a_list))
```

<!-- grade_id: cell-5949ac9777599dbc, points: 1, grade: true, solution: false -->
```python
assert_equal(has_duplicates([1, 2, 3]), False)
assert_equal(has_duplicates([1, 2, 2, 3]), True)
```

## Problem 4a
Create a class `Person` that has the following methods:
* `move` that moves the person's position by the optional parameters `dx` (default: 0) and `dy` (default: 0)
* `get_position` that returns the person's current position as a tuple of x and y coordinates
* `is_infected` that returns `True` if the person is infected and `False` otherwise
* `set_infected` that assigns the person as infected
* `get_infected` that accepts another `Person` and becomes infected by that person if that person is infected and their Euclidean distance is less than the second argument `threshold`

`Person` can also be optionally initialized by a tuple of x and y coordinates, indicating the person's initial position (default: (0,0)). A `Person` is initially uninfected.

<!-- grade_id: cell-7363386ed5790b38, points: 0, grade: false, solution: true -->
```python
import math
class Person:
    def __init__(self, initial_position=(0, 0)):
        self.x = initial_position[0]
        self.y = initial_position[1]
        self.infected = False

    def move(self, dx=0, dy=0):
        self.x += dx
        self.y += dy

    def get_position(self):
        return (self.x, self.y)

    def is_infected(self):
        return self.infected

    def set_infected(self):
        self.infected = True

    def get_infected(self, other_person, threshold):
        if other_person.is_infected():
            # Calculate Euclidean distance
            dx = self.x - other_person.x
            dy = self.y - other_person.y
            distance = math.sqrt(dx**2 + dy**2)
            if distance < threshold:
                self.set_infected()
```

<!-- grade_id: cell-07611009a85e2f77, points: 1, grade: true, solution: false -->
```python
person = Person()
assert_equal(person.get_position(), (0, 0))
assert_equal(person.is_infected(), False)

person.move(-1, 2)
assert_equal(person.get_position(), (-1, 2))
assert_equal(person.is_infected(), False)
```

<!-- grade_id: cell-7ee343c5ce1fa2b0, points: 1, grade: true, solution: false -->
```python
person1 = Person()
person2 = Person((1, 2))
person3 = Person((10, 20))
person3.set_infected()
person4 = Person((-1, -2))
person4.set_infected()

assert_equal(person1.is_infected(), False)
person1.get_infected(person2, 5)
assert_equal(person1.is_infected(), False)
person1.get_infected(person3, 5)
assert_equal(person1.is_infected(), False)
```

## Problem 4b
Create a class `QuarantinedPerson` that is a `Person` but whose position never changes regardless of the parameters passed to `move`.

<!-- grade_id: cell-41ba9afb8e1275e8, points: 0, grade: false, solution: true -->
```python
class QuarantinedPerson(Person):
    def __init__(self, initial_position=(0, 0)):
        super().__init__(initial_position)

    def move(self, dx=0, dy=0):
        # A quarantined person's position never changes
        pass
```

<!-- grade_id: cell-9a14040ff1676251, points: 1, grade: true, solution: false -->
```python
person1 = Person()
person2 = QuarantinedPerson((10, 0))
person2.set_infected()

assert_equal(person1.is_infected(), False)
person1.get_infected(person2, 5)
assert_equal(person1.is_infected(), False)
```

## Problem 5
Create a function `ten_div` that accepts a numerator and a denominator, and returns their quotient. The numerator can only be from 0 to 10, otherwise, it should raise a `TenDivError` which is a `ValueError`. It should return a `TenDivError` with message equal to `Error encountered: ` follow by the exception string if it encounters other exceptions.

<!-- grade_id: cell-804bdb56a0e766d8, points: 0, grade: false, solution: true -->
```python
class TenDivError(ValueError):
    pass

def ten_div(num, den):
    if not (0 <= num <= 10):
        raise TenDivError("Numerator must be between 0 and 10")
    try:
        return num / den
    except Exception as e:
        raise TenDivError(f"Error encountered: {e}")
```

<!-- grade_id: cell-503e96dc6ad1a424, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(ten_div(10, 2), 5.0)
assert_raises(TenDivError, ten_div, -2, 10)
assert_raises(TenDivError, ten_div, 11, 5)
```

## Problem 6
Create a function `z_per_element` that accepts the 2D `ndarray`s `x` and `y` and returns another 2D `ndarray` `z` defined by $z_{ij} = \exp(x_{ij}^2 + \cos(y_{ij})) + 2$.

```python
def z_per_element(x, y):
    """
    Calculates a 2D ndarray `z` based on the formula z_ij = exp(x_ij^2 + cos(y_ij)) + 2.

    Parameters
    ----------
    x : numpy.ndarray
        A 2D numpy array.
    y : numpy.ndarray
        A 2D numpy array with the same shape as `x`.

    Returns
    -------
    numpy.ndarray
        A 2D numpy array `z` computed element-wise.
    """
    z = np.exp(np.square(x) + np.cos(y)) + 2
    return z
```

```python
from numpy.testing import assert_array_equal
x = np.random.random((10, 2))
y = np.random.random((10, 2))
z = z_per_element(x, y)
assert_array_equal(z.shape, (10, 2))
```

## Problem 7
Create a function `row_dot` that accepts the 2D `ndarray`s `x` and `y` and returns another 1D `ndarray` `z` defined by $$z_i = \sum_j \sum_k x_{ik}y_{kj}.$$


```python
import numpy as np
def row_dot(x, y):
    """
    Calculates a 1D ndarray `z` based on the formula z_i = sum_j sum_k x_ik * y_kj.

    Parameters
    ----------
    x : numpy.ndarray
        A 2D numpy array.
    y : numpy.ndarray
        A 2D numpy array compatible for multiplication with `x`.

    Returns
    -------
    numpy.ndarray
        A 1D numpy array `z`.
    """
    # Perform matrix multiplication x @ y
    product_matrix = np.dot(x, y)
    # Sum across columns (axis=1) for each row
    z = np.sum(product_matrix, axis=1)
    return z
```

```python
x = np.random.random((10, 2))
y = np.random.random((2, 5))
z = row_dot(x, y)
assert_equal(z.shape, (10,))
```

## Problem 8
Create a function `shrink` that accepts a 2D `ndarray` `x` and returns another 2D `ndarray` `z` defined by
$$
z_{\left\lfloor i/2 \right\rfloor, j/2} = x_{ji},~~ i~\text{odd}, j~\text{even}
$$
$i$ and $j$ starts from zero and we consider zero to be even.

```python
def shrink(x):
    """
    Accepts a 2D ndarray `x` and returns another 2D ndarray `z` defined by
    z[floor(i/2), j/2] = x[j, i], where i is odd and j is even.

    Parameters
    ----------
    x : numpy.ndarray
        A 2D numpy array.

    Returns
    -------
    numpy.ndarray
        A 2D numpy array `z` with elements selected and rearranged from `x`.
    """
    # Select rows with even indices and columns with odd indices from x
    # x[j, i] where j is even, i is odd
    selected_elements = x[::2, 1::2]

    # The formula z[floor(i/2), j/2] = x[j, i]
    # implies that the first index of z (z_row) corresponds to original odd column (i) of x,
    # and the second index of z (z_col) corresponds to original even row (j) of x.
    # This is equivalent to transposing the `selected_elements` sub-array.
    z = selected_elements.T
    return z
```

```python
x = np.random.random((10, 8))
z = shrink(x)
assert_equal(z.shape, (4, 5))
```

## Problem 9
Create a function `multiplier` that takes in a 2D `ndarray` `x` and 1D array `y` and returns a 2D `ndarray` `z` given by $z_{ij} = x_{ij} * y_i$.

```python
def multiplier(x, y):
    """
    Calculates a 2D ndarray `z` where z_ij = x_ij * y_i.

    Parameters
    ----------
    x : numpy.ndarray
        A 2D numpy array.
    y : numpy.ndarray
        A 1D numpy array.

    Returns
    -------
    numpy.ndarray
        A 2D numpy array `z` with elements computed by multiplying
        each row of `x` by the corresponding element of `y`.
    """
    # Reshape y to (len(y), 1) to enable broadcasting across columns of x
    z = x * y[:, np.newaxis]
    return z
```

```python
x = np.random.random((10, 8))
y = np.random.random(10)
z = multiplier(x, y)
assert_equal(z.shape, (10, 8))
```

## Problem 10
Create a function `double_quadrant` that accepts a 2D ndarray `x` and returns another 2D `ndarray` `z` that has the same values as `x` except for the values in the upper left quadrant which are doubled. $x$ should remain the same.

```python
def double_quadrant(x):
    """
    Accepts a 2D ndarray `x` and returns another 2D ndarray `z`
    that has the same values as `x` except for the values in the
    upper left quadrant which are doubled. `x` should remain the same.

    Parameters
    ----------
    x : numpy.ndarray
        A 2D numpy array.

    Returns
    -------
    numpy.ndarray
        A 2D numpy array `z` with the upper left quadrant doubled.
    """
    # Create a copy of x to ensure x remains unchanged
    z = x.copy()

    # Get the dimensions of the array
    rows, cols = x.shape

    # Determine the boundaries of the upper left quadrant
    # Using integer division to handle odd dimensions gracefully
    mid_row = rows // 2
    mid_col = cols // 2

    # Double the values in the upper left quadrant of z
    z[0:mid_row, 0:mid_col] *= 2

    return z
```

```python
x = np.random.random((10, 8))
z = double_quadrant(x)
assert_equal(z.shape,(10, 8))
```