<a href="https://colab.research.google.com/github/ramiltrinidad-ctrl/maida/blob/main/Exercise_01RamilTrinidad.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
from numpy.testing import assert_equal
```

## Problem 1
Write a function `digit_sum` that takes a positive integer and returns the sum of its digits.

<!-- grade_id: cell-7169ef1194c4cc36, points: 0, grade: false, solution: true -->
```python
def digit_sum(number):
    ### YOUR CODE HERE

    total = 0
    while number > 0:
        total += number % 10
        number //= 10
    return total

test_value = int(input("Enter a number: "))
print(digit_sum(test_value))


```

Output:
```
Enter a number: 888
24
```

<!-- grade_id: cell-90ea299fbafa14e1, points: 1, grade: true, solution: false -->
```python
assert_equal(digit_sum(0), 0)
assert_equal(digit_sum(2020), 4)
assert_equal(digit_sum(1234), 10)
```

## Problem 2
Write a function `count_vowels` that returns the number of vowels in an input string.

<!-- grade_id: cell-661117ba16689550, points: 0, grade: false, solution: true -->
```python
def count_vowels(text):
    ### YOUR CODE HERE
    raise NotImplementedError
```

```python

```

<!-- grade_id: cell-661117ba16689550, points: 0, grade: false, solution: true -->
```python
def count_vowels(text):
    vowels = "aeiouAEIOU"
    count = 0
    for char in text:
        if char in vowels:
            count += 1
    return count

print (count_vowels("Hello World"))
```

Output:
```
3
```

<!-- grade_id: cell-9fe95b22b4da11bb, points: 1, grade: true, solution: false -->
```python
assert_equal(count_vowels('foo bar baz'), 4)
assert_equal(count_vowels('fbrbz'), 0)
```

## Problem 3
One classic method for composing secret messages is called a square code.  The spaces are removed from the english text and the characters are written into a square (or rectangle).  For example, the sentence "If man was meant to stay on the ground god would have given us roots" is 54 characters long, so it is written into a rectangle with 7 rows and 8 columns.

    ifmanwas
    meanttos        
    tayonthe
    groundgo
    dwouldha
    vegivenu
    sroots
                
The coded message is obtained by reading down the columns going left to right.   For example, the message above is coded as:

    imtgdvs fearwer mayoogo anouuio ntnnlvt wttddes aohghn sseoau

Write a function `encode` that takes in a a string message with no spaces between the words and returns the encode message.

```python
import math

def encode(text):
    length = len(text)
    if length == 0:
        return ""

    # Calculate dimensions for the square/rectangle
    columns = int(math.ceil(math.sqrt(length)))
    rows = int(math.ceil(length / float(columns)))

    encoded_message_parts = []
    for col in range(columns):
        current_column_chars = []
        for row in range(rows):
            index = row * columns + col
            if index < length:
                current_column_chars.append(text[index])
        encoded_message_parts.append("".join(current_column_chars))

    return " ".join(encoded_message_parts)

print(encode("MyNameisRamilTrinidadandIhaveasecretmessagewhichismillions"))
```

Output:
```
MRnIcain yaihrgss Nmdaeem aiavtwi mldemhl eTaaeil irnssci sidesho
```

<!-- grade_id: cell-0970aab1d2b5724f, points: 1, grade: true, solution: false -->
```python
assert_equal(encode('haveaniceday'), 'hae and via ecy')
assert_equal(encode('feedthedog'), 'fto ehg ee dd')
```

## Problem 4

Write a function `count_people` that accepts a multiline string of log entries. Each log entry consists of the string `IN` or `OUT` corresponding to whether the entry is the number of people who went in or out of the building, respectively, and the number of people this entry applies to, separate by the tab character (`\t`). `count_people` returns the number of people inside the building based on the log.

<!-- grade_id: cell-b3d234f4b1f5219e, points: 0, grade: false, solution: true -->
```python
def count_people(log):
    total_people = 0
    entries = log.strip().split('\n')
    for entry in entries:
        if entry:
            direction, num_str = entry.split('\t')
            num_people = int(num_str)
            if direction == 'IN':
                total_people += num_people
            elif direction == 'OUT':
                total_people -= num_people
    return total_people

# Test usage:
log_entries = """
IN	10
OUT	3
IN	5
OUT	2
IN	1
"""
print(count_people(log_entries))
```

Output:
```
11
```

<!-- grade_id: cell-0853c03ce6a292ba, points: 1, grade: true, solution: false -->
```python
assert_equal(
    count_people(
        'IN\t1\n'
        'OUT\t1\n'
        'IN\t1\n'
        'IN\t1\n'
        'IN\t1\n'
        'IN\t1\n'
        'OUT\t2\n'
    ),
    2
)
```

## Problem 5

Create a function `plot_bars` that accepts the string `heights` and returns a string which draws a plot based on `heights` when `print`ed. The plot should be 11 rows high, bordered at the sides by `|` and at the bottom by `+`. Each digit in `heights` corresponds to the number of `*` to draw from the bottom for that corresponding position.

<!-- grade_id: cell-cd4c0ef44b10a341, points: 0, grade: false, solution: true -->
```python
def plot_bars(heights):
    num_data_points = len(heights)

    # Initialize 10 rows for the stars, each filled with spaces.
    # The plot is 11 rows high: 10 for stars, 1 for the bottom border.
    # plot_content[0] will be the top-most row of stars, plot_content[9] the bottom-most.
    plot_content = [[' ' for _ in range(num_data_points)] for _ in range(10)]

    # Fill the plot_content with stars based on the heights string
    for col_idx, height_char in enumerate(heights):
        height = int(height_char)
        # Place stars from the bottom up for the current column
        for star_count in range(height):
            # Calculate the row index for the star (0 is top, 9 is bottom for stars)
            row_idx = 9 - star_count
            plot_content[row_idx][col_idx] = '*'

    # Build the final output string, adding side borders for each content row
    result_lines = []
    for row in plot_content:
        result_lines.append('|' + ''.join(row) + '|')

    # Add the bottom border
    result_lines.append('+' * (num_data_points + 2))

    return '\n'.join(result_lines)

# Test it
print(plot_bars('87654321'))
```

Output:
```
|        |
|        |
|*       |
|**      |
|***     |
|****    |
|*****   |
|******  |
|******* |
|********|
++++++++++
```

<!-- grade_id: cell-1460030c62c43423, points: 1, grade: true, solution: false -->
```python
assert_equal(
    plot_bars('12345678987654321'),
    '|                 |\n'
    '|        *        |\n'
    '|       ***       |\n'
    '|      *****      |\n'
    '|     *******     |\n'
    '|    *********    |\n'
    '|   ***********   |\n'
    '|  *************  |\n'
    '| *************** |\n'
    '|*****************|\n'
    '|+++++++++++++++++|\n'
)
assert_equal(
    plot_bars('98765432123456789'),
    '|                 |\n'
    '|*               *|\n'
    '|**             **|\n'
    '|***           ***|\n'
    '|****         ****|\n'
    '|*****       *****|\n'
    '|******     ******|\n'
    '|*******   *******|\n'
    '|******** ********|\n'
    '|*****************|\n'
    '|+++++++++++++++++|\n'
)
assert_equal(
    plot_bars('54321234543212345'),
    '|                 |\n'
    '|                 |\n'
    '|                 |\n'
    '|                 |\n'
    '|                 |\n'
    '|*       *       *|\n'
    '|**     ***     **|\n'
    '|***   *****   ***|\n'
    '|**** ******* ****|\n'
    '|*****************|\n'
    '|+++++++++++++++++|\n'
)
```