<a href="https://colab.research.google.com/github/ramiltrinidad-ctrl/maida/blob/main/RamilTrinidad_Exercise_03.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
from numpy.testing import assert_equal
```

For all problems, make sure to put in a complete docstring if there is none yet. You may only use the Python Standard Library.

## Problem 1
Write a function called `chop` that takes a list, modifies it by removing the first and last elements, and returns `None`.

```python
def chop(a_list):
    """
    Modifies a list by removing its first and last elements.

    Parameters
    ----------
    a_list : list
        The list to be modified.

    Returns
    -------
    None
    """
    if len(a_list) >= 2:
        del a_list[-1]  # Remove the last element
        del a_list[0]   # Remove the first element (which was the second originally)
    return None

## Testing
my_list = [1, 2, 3, 4, 5]
print(f"Original list: {my_list}")
chop(my_list)
print(f"List after chop: {my_list}")

# Test with a list of length less than 2
short_list = [100]
print(f"\nOriginal short list: {short_list}")
chop(short_list)
print(f"Short list after chop: {short_list}")

empty_list = []
print(f"\nOriginal empty list: {empty_list}")
chop(empty_list)
print(f"Empty list after chop: {empty_list}")

```

Output:
```
Original list: [1, 2, 3, 4, 5]
List after chop: [2, 3, 4]

Original short list: [100]
Short list after chop: [100]

Original empty list: []
Empty list after chop: []
```

<!-- grade_id: cell-645111f6cb08b363, points: 1, grade: true, solution: false -->
```python
a = [1, 2, 3, 4, 5]
assert chop(a) is None
assert_equal(a, [2, 3, 4])
```

## Problem 2
Write a function `sum_multiples` that computes the sum of multiples of 3 or 5 in an input list.

<!-- grade_id: cell-68e647afb7616bc5, points: 0, grade: false, solution: true -->
```python
def sum_multiples(a_list):

    total_sum = 0
    for num in a_list:
        if num % 3 == 0 or num % 5 == 0:
            total_sum += num
    return total_sum

# Input box for testing sum_multiples
user_input = input("Enter a comma-separated list of numbers (e.g., 1,2,3,4): ")

try:
    # Convert the input string to a list of integers
    user_list = [int(x.strip()) for x in user_input.split(',') if x.strip()]
    print(f"User list: {user_list}, Sum of multiples of 3 or 5: {sum_multiples(user_list)}")
except ValueError:
    print("Invalid input. Please enter numbers separated by commas.")

```

Output:
```
Enter a comma-separated list of numbers (e.g., 1,2,3,4): 0,0,0,0
User list: [0, 0, 0, 0], Sum of multiples of 3 or 5: 0
```

<!-- grade_id: cell-ba33eaa3bd8ee05a, points: 1, grade: true, solution: false -->
```python
assert_equal(sum_multiples(list(range(10))), 23)
assert_equal(sum_multiples(list(range(1, 10, 2))), 17)
assert_equal(sum_multiples(list(range(1, 10, 3))), 0)
```

## Problem 3
Write a function `rotate` that rotates a list by `k` elements.

<!-- grade_id: cell-7be8f562c1868e66, points: 0, grade: false, solution: true -->
```python
def rotate(numbers, k):
    """Rotate numbers by k elements"""
    ### YOUR CODE HERE
    if k > 0:
        return numbers[-k:] + numbers[:-k]
    elif k < 0:
        return numbers[-k:] + numbers[:-k]
    else:
        return numbers

```

<!-- grade_id: cell-8a8f9e9053ad1d3c, points: 1, grade: true, solution: false -->
```python
assert_equal(rotate([1, 2, 3, 4, 5, 6], 4), [3, 4, 5, 6, 1, 2])
assert_equal(rotate([1, 2, 3, 4, 5, 6], -4), [5, 6, 1, 2, 3, 4])
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_591/4103618171.py in <cell line: 0>()
----> 1 assert_equal(rotate([1, 2, 3, 4, 5, 6], 4), [3, 4, 5, 6, 1, 2])
      2 assert_equal(rotate([1, 2, 3, 4, 5, 6], -4), [5, 6, 1, 2, 3, 4])

NameError: name 'assert_equal' is not defined
```

## Problem 4
Write a function `on_all` that applies a function to every element of a list.

<!-- grade_id: cell-8205a5864316caaf, points: 0, grade: false, solution: true -->
```python
def on_all(func, a_list):
    """
    Applies a given function to every element of a list.

    Parameters
    ----------
    func : function
        The function to be applied to each element.
    a_list : list
        The list whose elements will be processed.

    Returns
    -------
    list
        A new list containing the results of applying `func` to each element.
    """
    result = []
    for item in a_list:
        result.append(func(item))
    return result
```

<!-- grade_id: cell-9c4be7904b0af281, points: 1, grade: true, solution: false -->
```python
assert_equal(on_all(abs, [1, -2, 3, -4, 5]), [1, 2, 3, 4, 5])
assert_equal(on_all(lambda x: x**2, [1, -2, 3, -4, 5]),
                   [1, 4, 9, 16, 25])
```

## Problem 5
Write a function `matrix_times_vector` that accepts matrix `mat` as a list of lists and vector `vec` as a list, and returns their product as a list. Assume the input dimensions are compatible.

<!-- grade_id: cell-51f7405f0170c1dd, points: 0, grade: false, solution: true -->
```python
def matrix_times_vector(mat, vec):

    result_vector = []
    for row in mat:
        dot_product = 0
        for i in range(len(row)):
            dot_product += row[i] * vec[i]
        result_vector.append(dot_product)
    return result_vector
```

<!-- grade_id: cell-679aba3b9570c7c0, points: 1, grade: true, solution: false -->
```python
assert_equal(
    matrix_times_vector(
        [[1, 2, 3],
         [7, 6, 8]],
        [9, 10, 11]
    ),
    [62, 211]
)
```

## Problem 6
Write a function `coder` that automatically converts alphanumeric English text to Morse code and vice versa. Empty `text` should return an empty string. Assume the message consists only of alphanumeric characters and space.

<!-- grade_id: cell-7cd6dfd0b24acd83, points: 0, grade: false, solution: true -->
```python
MORSE_CODE_DICT = {
    'A': '.-', 'B': '-...', 'C': '-.-.', 'D': '-..', 'E': '.', 'F': '..-.',
    'G': '--.', 'H': '....', 'I': '..', 'J': '.---', 'K': '-.-', 'L': '.-..',
    'M': '--', 'N': '-.', 'O': '---', 'P': '.--.', 'Q': '--.-', 'R': '.-.',
    'S': '...', 'T': '-', 'U': '..-', 'V': '...-', 'W': '.--', 'X': '-..-',
    'Y': '-.--', 'Z': '--..',
    '0': '-----', '1': '.----', '2': '..---', '3': '...--', '4': '....-',
    '5': '.....', '6': '-....', '7': '--...', '8': '---..', '9': '----.',
    ' ': '   '
}

# Reverse dictionary for decoding Morse to English
DECODE_MORSE_DICT = {value: key for key, value in MORSE_CODE_DICT.items()}
DECODE_MORSE_DICT[''] = '' # Handle empty string for decoding

def coder(text, to_morse=True):

    if not text:
        return ""

    if to_morse:
        coded_text = []
        for char in text.upper():
            coded_text.append(MORSE_CODE_DICT.get(char, '')) # .get handles chars not in dict
        return ' '.join(coded_text)
    else:
        # Split Morse code by '   ' for spaces between words
        words_in_morse = text.split('   ')
        decoded_words = []
        for word_morse in words_in_morse:
            # Split each word by ' ' for spaces between letters
            letters_in_morse = word_morse.split(' ')
            decoded_letters = []
            for letter_morse in letters_in_morse:
                decoded_letters.append(DECODE_MORSE_DICT.get(letter_morse, ''))
            decoded_words.append(''.join(decoded_letters))
        return ' '.join(decoded_words)
```

<!-- grade_id: cell-dc1c5996cd5902b2, points: 1, grade: true, solution: false -->
```python
assert_equal(
    coder('Hello world'),
    '.... . .-.. .-.. ---   .-- --- .-. .-.. -..'
)
assert_equal(
    coder('-.. .- - .-   ... -.-. .. . -. -.-. .', False),
    'DATA SCIENCE'
)
```

## Problem 7
Complete the function below.

<!-- grade_id: cell-ac1556c1573c42d3, points: 0, grade: false, solution: true -->
```python
def sort_by_key(items_with_keys, ascending=True):

    # Create a copy to ensure the original list remains unmodified
    sorted_list = sorted(items_with_keys, key=lambda x: (x[1], x[0]), reverse=not ascending)
    return sorted_list
```

<!-- grade_id: cell-da5dfd9915098c50, points: 1, grade: true, solution: false -->
```python
assert_equal(
    sort_by_key([('B', 10), ('A', 12), ('C', 3), ('D', 5)]),
    [('C', 3), ('D', 5), ('B', 10), ('A', 12)])
assert_equal(
    sort_by_key([('B', 10), ('A', 12), ('C', 3), ('D', 5)], False),
    [('A', 12), ('B', 10), ('D', 5), ('C', 3)])
```

## Problem 8
Write a function `count_words` that accepts a string `text` and returns a dictionary where the key is a word and the value is the number of occurrences of that word in `text`. A word is defined as a case-insensitive sequence of non-whitespace characters.

<!-- grade_id: cell-bc907ec0368f8aaf, points: 0, grade: false, solution: true -->
```python
def count_words(text):

    word_counts = {}
    # Convert to lowercase and split by whitespace
    words = text.lower().split()

    for word in words:
        # Remove punctuation from words (optional, but good for robust counting)
        # For this problem, 'non-whitespace characters' implies we don't need to strip punctuation explicitly,
        # but it's often a good practice for word counting.
        # Let's assume words are already clean or that punctuation attached to words should be considered part of the word
        # as per 'non-whitespace characters'.
        # Given the example 'lamb', 5 and 'it's', 1, it seems basic split works and punctuation should be part of the word.
        word_counts[word] = word_counts.get(word, 0) + 1

    return word_counts


```

<!-- grade_id: cell-6e8fc7aff97f3758, points: 1, grade: true, solution: false -->
```python
wc = count_words("""
Mary had a little lamb
Little lamb little lamb
Mary had a little lamb
It's fleece was white as snow

Everywhere that Mary went
Mary went Mary went
Everywhere that Mary went
The lamb was sure to go
""")
assert_equal(len(wc), 18)
assert_equal(wc['lamb'], 5)
assert_equal(wc["it's"], 1)
```

## Problem 9

Write a function `display_tree` that accepts a dictionary whose values can be numbers, strings or dictionaries, and returns a string. The contents of the returned string is as follows:
* For every key-value pair in the dictionary, sort by key then show the key, one per line, followed by a colon (:)
* If the value is not a dictionary, show the value prefixed by a space
* If the value is a dictionary, display the contents on the dictionary from the next line, following the rules in this list
* Indent by two spaces for every level of the dictionary
* Use the newline character (`\n`) for creating lines

<!-- grade_id: cell-233f3054b37827fc, points: 0, grade: false, solution: true -->
```python
def display_tree(data, indent_level=0):

    output_lines = []
    indent = '  ' * indent_level

    # Sort keys for consistent output
    sorted_keys = sorted(data.keys(), key=lambda x: str(x))

    for key in sorted_keys:
        value = data[key]
        if isinstance(value, dict):
            output_lines.append(f"{indent}{key}:")
            # Recursively call for nested dictionaries, increasing indent_level
            output_lines.append(display_tree(value, indent_level + 1))
        else:
            output_lines.append(f"{indent}{key}: {value}")

    # Join lines with newline characters and filter out empty strings from recursion
    result = '\n'.join(filter(None, output_lines))
    # Add a trailing newline if the result is not empty and it's the top-level call
    # This is to match the assert_equal's desired output format
    if result and indent_level == 0:
        result += '\n'
    return result
```

<!-- grade_id: cell-63304af010a78f81, points: 1, grade: true, solution: false -->
```python
assert_equal(
    display_tree({
    'foo': 1,
    'bar': {
        'fee': 'eenie',
        'fi': 'meenie',
        'fo': 'minnie',
        'fum': {
            'mo': 'lover'
        },
    'baz':
        {
            1: 'potato',
            2: 'potatoes',
            3: 'potatoes'
        }
    }
    }),
    """bar:
  baz:
    1: potato
    2: potatoes
    3: potatoes
  fee: eenie
  fi: meenie
  fo: minnie
  fum:
    mo: lover
foo: 1
"""
)
```

## Problem 10

Write a function `get_nested_key_value` that accepts a nested dictionary `nested_dict` and a string `key`, which corresponds to the key at each level joined by a period (.). Missing keys should return `None`. Assume that keys are alphanumeric.

<!-- grade_id: cell-276faaed0193b2c3, points: 0, grade: false, solution: true -->
```python
def get_nested_key_value(nested_dict, key):
    """
    Retrieves a value from a nested dictionary using a dot-separated key string.

    Parameters
    ----------
    nested_dict : dict
        The nested dictionary to search.
    key : str
        A string representing the path to the desired value, with keys separated by periods (e.g., 'a.b.c').

    Returns
    -------
    any or None
        The value associated with the key path, or None if any key in the path is missing.
    """
    keys = key.split('.')
    current_dict = nested_dict

    for k in keys:
        if isinstance(current_dict, dict) and k in current_dict:
            current_dict = current_dict[k]
        else:
            return None
    return current_dict
```

<!-- grade_id: cell-99592c2742ae0394, points: 1, grade: true, solution: false -->
```python
d = {
    'a': {
        'a1': {
            'aa1': 'bar',
            'aa2': 'baz'
        },
        'a2': 'bazbar',
        'a3': 0
    },
    'b': 1,
    'c': {
        'c1': 'eenie',
        'c2': {
            'x': 'meenie',
            'y': 'minnie',
            'z': 'mo'
        },
        'c3': 'lover'
    }
}
assert_equal(
    get_nested_key_value(d, 'a.a1.aa2'),
    'baz'
)
assert_equal(
    get_nested_key_value(d, 'a.a1'),
    {
        'aa1': 'bar',
        'aa2': 'baz'
    }
)
assert_equal(
    get_nested_key_value(d, 'a.a1.aa2.aaa1'),
    None
)
assert_equal(
    get_nested_key_value(d, 'a.a1.aa5'),
    None
)
```