```python
from numpy.testing import assert_equal, assert_allclose
```

## Problem 1
Write a Python function that will count from 1 to an integer `n` (inclusive). For each count, print the counter value and a message separated by a single space. If the counter is even, the message should be `fizz`, `fuzz` if divisible by 3 and `foo` (only) if divisible by 6.

<!-- grade_id: cell-0e7266cd03645e41, points: 0, grade: false, solution: true -->
```python
def count_and_print(n):
    """Count from 1 to n whilst displaying a message

    Count from 1 to an integer n (inclusive). For each count, print the
    counter value and a message separated by a single space. If the counter is
    even, the message should be `fizz`, `fuzz` if divisible by 3 and `foo`
    (only) if divisible by 6.

    Parameters
    ----------
    n : integer
        Count up to this number (inclusive)

    Examples
    --------
    >>> count_and_print(12)
    1
    2 fizz
    3 fuzz
    4 fizz
    5
    6 foo
    7
    8 fizz
    9 fuzz
    10 fizz
    11
    12 foo

    """
    for i in range(1, n + 1):
        message = []
        if i % 6 == 0:
            message.append('foo')
        elif i % 3 == 0:
            message.append('fuzz')
        elif i % 2 == 0:
            message.append('fizz')

        if message:
            print(f"{i} {' '.join(message)}")
        else:
            print(f"{i}")

## To test
count_and_print(12)

count_and_print(18)
```

Output:
```
1
2 fizz
3 fuzz
4 fizz
5
6 foo
7
8 fizz
9 fuzz
10 fizz
11
12 foo
1
2 fizz
3 fuzz
4 fizz
5
6 foo
7
8 fizz
9 fuzz
10 fizz
11
12 foo
13
14 fizz
15 fuzz
16 fizz
17
18 foo
```

```python
%%capture out1
count_and_print(12)
```

```python
%%capture out2
count_and_print(18)
```

<!-- grade_id: cell-5e68bd9f3673ce01, points: 1, grade: true, solution: false -->
```python
assert_equal(
    out1.stdout,
    '1\n'
    '2 fizz\n'
    '3 fuzz\n'
    '4 fizz\n'
    '5\n'
    '6 foo\n'
    '7\n'
    '8 fizz\n'
    '9 fuzz\n'
    '10 fizz\n'
    '11\n'
    '12 foo\n')
assert_equal(
    out2.stdout,
    '1\n'
    '2 fizz\n'
    '3 fuzz\n'
    '4 fizz\n'
    '5\n'
    '6 foo\n'
    '7\n'
    '8 fizz\n'
    '9 fuzz\n'
    '10 fizz\n'
    '11\n'
    '12 foo\n'
    '13\n'
    '14 fizz\n'
    '15 fuzz\n'
    '16 fizz\n'
    '17\n'
    '18 foo\n'
)
```

## Problem 2
Write a function `overlap_interval` that accepts parameters `start1`, `end1`, `start2` and `end2` corresponding to the endpoints of two numeric intervals. It should print the endpoints of their overlaps or print the message `No overlap!` if they don't overlap.

<!-- grade_id: cell-c3a64bdf8b11fb95, points: 0, grade: false, solution: true -->
```python
def overlap_interval(start1, end1, start2, end2):
    # Calculate the start and end points of the potential overlap
    overlap_start = max(start1, start2)
    overlap_end = min(end1, end2)

    # Check if there is an actual overlap
    if overlap_start < overlap_end:
        print(f"{overlap_start} {overlap_end}")
    else:
        print("No overlap!")

overlap_interval(-5, 10, 5, 12)
overlap_interval(-5, 10, 11, 12)
overlap_interval(-5, 10, -10, 12)
overlap_interval(5, 8, -10, 12)
overlap_interval(-12, 12, 8, 10)
```

Output:
```
5 10
No overlap!
-5 10
5 8
8 10
```

```python
%%capture out1
overlap_interval(-5, 10, 5, 12)
```

```python
%%capture out2
overlap_interval(-5, 10, 11, 12)
```

```python
%%capture out3
overlap_interval(-5, 10, -10, 12)
```

```python
%%capture out4
overlap_interval(5, 8, -10, 12)
```

```python
%%capture out5
overlap_interval(-12, 12, 8, 10)
```

<!-- grade_id: cell-7812ed25101d9ddf, points: 1, grade: true, solution: false -->
```python
assert_equal(out1.stdout, '5 10\n')
assert_equal(out2.stdout, 'No overlap!\n')
assert_equal(out3.stdout, '-5 10\n')
```

## Problem 3
Write a function that returns the cosine of an input (in radians) using [Maclaurin's approximation](https://en.wikibooks.org/wiki/Trigonometry/Power_Series_for_Cosine_and_Sine). It should use a `while` loop to compute terms of the summation until the last term is smaller than $10^{-15}$.

<!-- grade_id: cell-54183911ca9ecc32, points: 0, grade: false, solution: true -->
```python
import math

def cosine(theta):
    """Return cosine of theta using Maclaurin's approximation

    Parameters
    ----------
    theta : float
        Value to compute cosine of, in radians

    Returns
    -------
    cos_theta : float
        Cosine of theta
    """
    cos_theta = 0.0
    n = 0
    tolerance = 1e-15

    while True:
        term = ((-1)**n) * (theta**(2*n)) / math.factorial(2*n)
        print(f"Term {n}: {term}")

        # Stop if the absolute value of the current term is less than tolerance,
        # but ensure the first term (n=0) is always added.
        if n > 0 and abs(term) < tolerance:
            break

        cos_theta += term
        n += 1

    return cos_theta

##cosine(0)
##cosine(math.pi/2)
cosine(math.pi/4)
##cosine(math.pi/6)
```

Output:
```
Term 0: 1.0
Term 1: -0.30842513753404244
Term 2: 0.0158543442438155
Term 3: -0.00032599188692738996
Term 4: 3.590860448591509e-06
Term 5: -2.461136950494199e-08
Term 6: 1.1501159127974046e-10
Term 7: -3.8980731712596733e-13
Term 8: 1.0018864616362713e-15
Term 9: -2.0196533968866804e-18
0.7071067811865475
```

<!-- grade_id: cell-356698de50ce75c6, points: 1, grade: true, solution: false -->
```python
from math import pi, sqrt
atol = 1e-15 # tolerance
assert_allclose(cosine(0), 1.0, rtol=0, atol=atol)
assert_allclose(cosine(pi/2) , 0, rtol=0, atol=atol)
assert_allclose(cosine(pi/4) , sqrt(2)/2, rtol=0, atol=atol)
assert_allclose(cosine(pi/6) , sqrt(3)/2, rtol=0, atol=atol)
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_1463/1069226291.py in <cell line: 0>()
      1 from math import pi, sqrt
      2 atol = 1e-15 # tolerance
----> 3 assert_allclose(cosine(0), 1.0, rtol=0, atol=atol)
      4 assert_allclose(cosine(pi/2) , 0, rtol=0, atol=atol)
      5 assert_allclose(cosine(pi/4) , sqrt(2)/2, rtol=0, atol=atol)

NameError: name 'assert_allclose' is not defined
```

## Problem 4
The greatest common divisor (GCD) of $a$ and $b$ is the largest number that divides both of them with no remainder.

One way to find the GCD of two numbers is based on the observation that if $r$ is the remainder when $a$ is divided by $b$, then $gcd(a, b) = gcd(b, r)$. As a base case, we can use $gcd(a, 0) = a$.

Write a function called `gcd` that takes parameters `a` and `b` and returns their greatest common divisor.

<!-- grade_id: cell-c29ca12ec75789d2, points: 0, grade: false, solution: true -->
```python
def gcd(a, b):
    if b == 0:
        return a
    else:
        return gcd(b, a % b)

a = int(input("Enter first number (a): "))
b = int(input("Enter second number (b): "))

x = gcd(a, b)

print(f"The GCD of {a} and {b} is: {x}")
```

Output:
```
Enter first number (a): 13
Enter second number (b): 5
The GCD of 13 and 5 is: 1
```

<!-- grade_id: cell-2f070dd683636894, points: 1, grade: true, solution: false -->
```python
assert_equal(gcd(4, 2), 2)
assert_equal(gcd(2, 4), 2)
assert_equal(gcd(13, 5), 1)
assert_equal(gcd(20, 24), 4)
```

## Problem 5
Write a function `biased_sum` that returns the sum of positive integers passed to it but numbers that are multiples of the `bias` are doubled before being added. By default, `bias` is 2. Return `None` if there are no numbers passed to it. Ignore arguments that are not positive integers.

```python
def biased_sum(*args, bias=2):
    if not args:
        return None
    total = 0
    for arg in args:
        if isinstance(arg, int) and arg > 0:
            if arg % bias == 0:
                total += arg * 2
            else:
                total += arg
    return total

results = []
results.append(biased_sum(1, 2, 3))
results.append(biased_sum(1, 2, 3, 4, 5, 6, 3))
results.append(biased_sum(1, 2, 3, 4, 5, 6, bias=3))
results.append(biased_sum(1, 2, 3, 4, 5, -5, 6, bias=3))

print("List all results:")
for r in results:
    print(r)
```

Output:
```
List all results:
8
36
30
30
```

<!-- grade_id: cell-269f051cdd06af87, points: 1, grade: true, solution: false -->
```python
assert_equal(biased_sum(1, 2, 3), 8)
assert_equal(biased_sum(1, 2, 3, 4, 5, 6, 3), 36)
assert_equal(biased_sum(1, 2, 3, 4, 5, 6, bias=3), 30)
assert_equal(biased_sum(1, 2, 3, 4, 5, -5, 6, bias=3), 30)
```

## Problem 6
Write a function `last_in_sequence` that accepts a string of digits, which can have a repeating sequence of 0 to 9 in it but may have intervening digits between them. The function should return the last digit in the sequence found in the string. It should return `None` if there is no such sequence in the string.

Examples:
* **0**9**12**251**3**2**4**21 $\rightarrow$ 4
* 1**0**9**12**251**3**2**4**21 $\rightarrow$ 4
* **012345678901234** $\rightarrow$ 4
* 1912251324 $\rightarrow$ `None`

<!-- grade_id: cell-f26b1b122b693409, points: 0, grade: false, solution: true -->
```python
def last_in_sequence(s):
    last_found_digit = None
    current_expected_digit = 0 # We always start looking for 0

    for char in s:
        if not char.isdigit():
            continue # Ignore non-digit characters

        digit = int(char)

        if digit == current_expected_digit:
            last_found_digit = digit
            current_expected_digit = (current_expected_digit + 1) % 10

    return last_found_digit

print(last_in_sequence('0123456789012345678'))
print(last_in_sequence('0912251324'))
print(last_in_sequence('10912251324'))
print(last_in_sequence('1912251324'))

```

Output:
```
8
4
4
None
```

<!-- grade_id: cell-19149d85bde88dea, points: 1, grade: true, solution: false -->
```python
assert_equal(last_in_sequence('0912251324'), 4)
assert_equal(last_in_sequence('10912251324'), 4)
assert_equal(last_in_sequence('012345678901234'), 4)
assert_equal(last_in_sequence('1912251324'), None)
```

## Problem 7
Write a function `check_password` that returns `True` if the input password satisfies the following criteria:

* at least 8 characters in length
* contains at least one lowercase letter
* contains at least one uppercase letter
* contains at least one number

It should return `False` if at least one of the criteria is not satisfied.

<!-- grade_id: cell-a72d390a2aa6048b, points: 0, grade: false, solution: true -->
```python
def check_password(password):
    # at least 8 characters in length
    if len(password) < 8:
        return False
    # contains at least one lowercase letter
    if not any(char.islower() for char in password):
        return False
    # contains at least one uppercase letter
    if not any(char.isupper() for char in password):
        return False
    # contains at least one number
    if not any(char.isdigit() for char in password):
        return False

    # All criteria met
    return True

password_input = input("Enter a password to check: ")
if check_password(password_input):
    print(f"The password '{password_input}' is True.")
else:
    print(f"The password '{password_input}' is False.")

```

Output:
```
Enter a password to check: 12345ABC
The password '12345ABC' is False.
```

<!-- grade_id: cell-b8041b3660eec0bd, points: 1, grade: true, solution: false -->
```python
assert_equal(check_password('abcABC1'), False)
assert_equal(check_password('abcABC12'), True)
assert_equal(check_password('abcABCDE'), False)
assert_equal(check_password('abc12345'), False)
assert_equal(check_password('12345ABC'), False)

```

## Problem 8
Write a function `is_palindrome` that returns `True` if the input string is a palindrome and `False` otherwise. In this case, an empty string is **NOT** a palindrome.

<!-- grade_id: cell-c7e575e3d909fd56, points: 0, grade: false, solution: true -->
```python
def is_palindrome(text):
    # An empty string is NOT a palindrome
    if not text:
        return False

    # Clean the string: remove spaces and convert to lowercase
    cleaned_text = "".join(char.lower() for char in text if char.isalnum())

    # Check if the cleaned string is equal to its reverse
    return cleaned_text == cleaned_text[::-1]

is_palindrome('abcd')
is_palindrome('ab ba')
```

Output:
```
True
```

<!-- grade_id: cell-71c90ba9148f6663, points: 1, grade: true, solution: false -->
```python
assert_equal(is_palindrome('abcd'), False)
assert_equal(is_palindrome('abba'), True)
assert_equal(is_palindrome('abeba'), True)
assert_equal(is_palindrome('abcda'), False)
assert_equal(is_palindrome('ab cd'), False)
assert_equal(is_palindrome('ab ba'), True)
assert_equal(is_palindrome('ab ab'), False)
assert_equal(is_palindrome('ab  ba'), True)
assert_equal(is_palindrome(' ab   ba     '), True)
```

## Problem 9
Create a function `print_patient_info` that accepts `name`, `age`, `sex`, `height`, `weight`, `icd_code` and `icd_desc` then prints them horizontally, in a columnar manner. It should be preceded by a line of headers centered in its column. Separate each column by a space. The details of the different columns are as follows:

| Variable name | Header          | Column width | Alignment | Formatting       |
|---------------|-----------------|--------------|-----------|------------------|
| `name`        | Name            | 20           | left      |                  |
| `age`         | Age             | 3            | right     |                  |
| `sex`         | Sex             | 3            | center    |                  |
| `height`      | Height          | 6            | right     | 2 decimal places |
| `weight`      | Weight          | 6            | right     | 2 decimal places |
| `icd_code`    | ICD Code        | 10           | center    |                  |
| `icd_desc`    | ICD Description | 50           | left      |                  |

<!-- grade_id: cell-42fb36e8cac45087, points: 0, grade: false, solution: true -->
```python
def print_patient_info(name, age, sex, height, weight, icd_code, icd_desc):
    # Define headers and their properties
    headers = {
        'name': ('Name', 20, '^'),
        'age': ('Age', 3, '^'),
        'sex': ('Sex', 3, '^'),
        'height': ('Height', 6, '^'),
        'weight': ('Weight', 6, '^'),
        'icd_code': ('ICD Code', 10, '^'),
        'icd_desc': ('ICD Description', 50, '^')
    }

    # Construct header line
    header_line = (
        f"{headers['name'][0]:{headers['name'][2]}{headers['name'][1]}}" + ' ' +
        f"{headers['age'][0]:{headers['age'][2]}{headers['age'][1]}}" + ' ' +
        f"{headers['sex'][0]:{headers['sex'][2]}{headers['sex'][1]}}" + ' ' +
        f"{headers['height'][0]:{headers['height'][2]}{headers['height'][1]}}" + ' ' +
        f"{headers['weight'][0]:{headers['weight'][2]}{headers['weight'][1]}}" + ' ' +
        f"{headers['icd_code'][0]:{headers['icd_code'][2]}{headers['icd_code'][1]}}" + ' ' +
        f"{headers['icd_desc'][0]:{headers['icd_desc'][2]}{headers['icd_desc'][1]}}"
    )

    # Construct data line
    data_line = (
        f"{name:<20}" + ' ' +
        f"{age:>3}" + ' ' +
        f"{sex:^3}" + ' ' +
        f"{height:>6.2f}" + ' ' +
        f"{weight:>6.2f}" + ' ' +
        f"{icd_code:^10}" + ' ' +
        f"{icd_desc:<50}"
    )

    print(header_line)
    print(data_line)

print_patient_info('J. De la Cruz', 34, 'M', 1.732, 112, 'E11.621',
                   'Type 2 diabetes mellitus with foot ulcer (E)')
```

Output:
```
Name         Age Sex Height Weight  ICD Code                   ICD Description                  
J. De la Cruz         34  M    1.73 112.00  E11.621   Type 2 diabetes mellitus with foot ulcer (E)
```

```python
%%capture out1
print_patient_info('J. De la Cruz', 34, 'M', 1.732, 112, 'E11.621',
                   'Type 2 diabetes mellitus with foot ulcer (E)')
```

```python
%%capture out2
print_patient_info('M. Clara', 16, 'F', 1.52, 48, 'Z79.891',
                   'Long term (current) use of opiate analgesic (O)')
```

```python
%%capture out3
print_patient_info('P. Tasyo', 64, 'M', 1.6, 60, 'Z47.1',
                   'Aftercare following joint replacement surgery')
```

<!-- grade_id: cell-c402616ea73275b4, points: 1, grade: true, solution: false -->
```python
assert_equal(
    out1.stdout,
    '        Name         Age Sex Height Weight  ICD Code                   '
    'ICD Description                  \n'
    'J. De la Cruz         34  M    1.73 112.00  E11.621   '
    'Type 2 diabetes mellitus with foot ulcer (E)      \n'
)
assert_equal(
    out2.stdout,
    '        Name         Age Sex Height Weight  ICD Code                   '
    'ICD Description                  \n'
    'M. Clara              16  F    1.52  48.00  Z79.891   '
    'Long term (current) use of opiate analgesic (O)   \n'
)
```

## Problem 10
Write a function `create_grid` that takes in the number of squares on a side and the number of stars on a side of a square then returns the grid of squares as a string. Separate each line by the newline (`\n`) character.


<!-- grade_id: cell-d554859dd48da4e8, points: 0, grade: false, solution: true -->
```python
def create_grid(num_squares, num_stars):
    grid_lines = []

    # Define the content for horizontal segments and square contents
    h_line_segment_content = ('- ' * num_stars).strip()
    star_line_segment_content = ('* ' * num_stars).strip()
    # Ensure empty segment has the same length as the star segment for alignment
    empty_line_segment_content = ' ' * len(star_line_segment_content)

    # Iterate through each row of squares
    for r_sq in range(num_squares):
        # Construct a horizontal border line (e.g., '+ - - + - - +')
        border_parts = ['+']
        for _ in range(num_squares):
            border_parts.append(h_line_segment_content)
            border_parts.append('+')
        grid_lines.append(' '.join(border_parts))

        # Iterate through each star row within the current square row
        for _ in range(num_stars):
            content_parts = ['|']
            # Iterate through each column of squares
            for c_sq in range(num_squares):
                # Determine if the current square should have stars or be empty
                # The pattern alternates: (0,0) stars, (0,1) empty, (1,0) empty, (1,1) stars
                if (r_sq + c_sq) % 2 == 0:
                    current_segment = star_line_segment_content
                else:
                    current_segment = empty_line_segment_content
                content_parts.append(current_segment)
                content_parts.append('|')
            grid_lines.append(' '.join(content_parts))

    # Add the final horizontal border line
    border_parts = ['+']
    for _ in range(num_squares):
        border_parts.append(h_line_segment_content)
        border_parts.append('+')
    grid_lines.append(' '.join(border_parts))

    return '\n'.join(grid_lines) + '\n'

create_grid(1, 1)
```

Output:
```
'+ - +\n| * |\n+ - +\n'
```

<!-- grade_id: cell-804b324f525cec69, points: 1, grade: true, solution: false -->
```python
assert_equal(
    create_grid(1, 1),
"""+ - +
| * |
+ - +
""")
assert_equal(
    create_grid(2, 4),
"""+ - - - - + - - - - +
| * * * * |         |
| * * * * |         |
| * * * * |         |
| * * * * |         |
+ - - - - + - - - - +
|         | * * * * |
|         | * * * * |
|         | * * * * |
|         | * * * * |
+ - - - - + - - - - +
""")
assert_equal(
    create_grid(4, 2),
"""+ - - + - - + - - + - - +
| * * |     | * * |     |
| * * |     | * * |     |
+ - - + - - + - - + - - +
|     | * * |     | * * |
|     | * * |     | * * |
+ - - + - - + - - + - - +
| * * |     | * * |     |
| * * |     | * * |     |
+ - - + - - + - - + - - +
|     | * * |     | * * |
|     | * * |     | * * |
+ - - + - - + - - + - - +
""")
```